

# Homework 1: Basic Python, Arrays, and DataFrames

## Due Thursday, October 8th at 11:59PM

Welcome to Homework 1! This week's homework will cover basic Python, arrays, and DataFrames. You can find additional help on these topics in [Chapter 1](https://www.inferentialthinking.com/chapters/01/what-is-data-science.html) of Computational and Inferential Thinking and [BPD 1-11](https://notes.dsc10.com/01-getting_started/tools.html) in the `babypandas` notes.


### Instructions

Remember to start early and submit often. You are given six slip days throughout the quarter to extend deadlines. See the syllabus for more details. With the exception of using slip days, late work will not be accepted unless you have made special arrangements with your instructor.

**Important**: For homeworks, the `otter` tests don't usually tell you that your answer is correct. More often, they help catch careless mistakes. It's up to you to ensure that your answer is correct. If you're not sure, ask someone (not for the answer, but for some guidance about your approach). These are great questions for office hours (the schedule can be found [here](https://dsc10.com/calendar)) or Campuswire. Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

```python
# Please don't change this cell, but do make sure to run it.
import babypandas as bpd
import matplotlib.pyplot as plt
import numpy as np
np.set_printoptions(legacy='1.21')
import otter
grader = otter.Notebook()

plt.style.use('ggplot')
```

## 1. The Three Musketeers vs. Les Trois Mousquetaires 🐭🐭🐭

<center><img src=images/3_musketeers_illustration.png width=350>(<a href="https://en.wikipedia.org/wiki/The_Three_Musketeers">source</a>)</center>

In Lecture 1, we counted the number of times that the characters Amy, Beth, Jo, Meg, and Laurie were named in each chapter of the classic book, Little Women. In programming, the word "character" also refers to a single element of a string. For instance, the string `"3 zebras!"` has 9 characters – `"3"`, `" "`, `"z"`, `"e"`, `"b"`, `"r"`, `"a"`, `"s"`, and `"!"`. 

Let's use this concept to see if *The Three Musketeers* by Alexandre Dumas has longer sentences in English or the original French. 

The following code generates a scatter plot where each dot represents a chapter of *The Three Musketeers*, either in English or French. For each chapter, there is both a blue star corresponding to that chapter's English translation, and a green triangle for the original French chapter. The horizontal position of a dot measures the number of periods in the chapter. The vertical position measures the total number of characters in that chapter. 

```python
# This cell contains code that hasn't yet been covered in the course.
# It isn't expected that you'll understand the code, but you should be able to 
# interpret the scatter plot it generates.

eng_file = "./data/The_Three_Musketeers.txt"
fre_file = "./data/Les_Trois_Mousquetaires.txt"

eng_chapters = open(eng_file, encoding="utf-8").read().split('Chapter ')[67:]
fre_chapters = open(fre_file, encoding="utf-8").read().split('CHAPITRE ')[67:]

eng_periods = np.char.count(eng_chapters, '.')
fre_periods = np.char.count(fre_chapters, '.')

eng_chars = [len(c) for c in eng_chapters]
fre_chars = [len(c) for c in fre_chapters]

plt.scatter(eng_periods, eng_chars, color='b', marker='*')
plt.scatter(fre_periods, fre_chars, color='g', marker='^')
plt.xlabel('Periods')
plt.ylabel('Characters')
plt.legend(['English', 'French'])
plt.axis([0, 450, 0, 45000]); 
```

Before we decide which language has longer sentences, let's start with some investigative questions. All of the questions in this section should be answered using the scatter plot alone, plus process of elimination and estimation. Don't write any code for this section! It may help to open another copy of the notebook in a new tab so you can easily refer to the scatter plot throughout this section.

**Question 1.1.** How many <ins>periods</ins> are in the **English** chapter with the greatest number of <ins>characters</ins>? Assign the variable `musketeers_1` to either 1, 2, 3, 4, or 5 corresponding to your choice.

1. 325
2. 352
3. 400
4. 405
5. 452

```python
musketeers_1 = ...
```

```python
grader.check("q1_1")
```

**Question 1.2.** How many <ins>characters</ins> are in the **French** chapter with the most <ins>periods</ins>? Assign the variable `musketeers_2` to either 1, 2, 3, 4, or 5.

1. 27,625
2. 32,625
3. 37,625
4. 39,625
5. 40,625

```python
musketeers_2 = ...
```

```python
grader.check("q1_2")
```

**Question 1.3.** Which of the following is closest to the <ins>average number of characters per period</ins> in the **English** version of *The Three Musketeers*? This is roughly the length of a typical sentence in the English translation of the book. Assign the variable `musketeers_3` to either 1, 2, 3, 4, or 5

1. 50
2. 100
3. 150
4. 200
5. 250

```python
musketeers_3 = ...
```

```python
grader.check("q1_3")
```

**Question 1.4.** Which version of *The Three Musketeers* has more characters per period, on average? Assign the variable `musketeers_4` to either 1 or 2.

1. French
2. English

```python
musketeers_4 = ...
```

```python
grader.check("q1_4")
```

**Question 1.5.** Which of the following is a valid conclusion we can make based off of the above scatter plot alone? Assign the variable `musketeers_5` to 1, 2, 3, 4, or 5. There is only one correct answer.

1. The number of periods in each chapter is the same in the French and English versions of the book.

1. The number of characters in each chapter is the same in the French and English versions of the book.

1. The last chapter has more periods in the English translation than in the original French.

1. In the first chapter, the average number of characters per period is roughly the same in the French and English versions of the book.

1. None of these is a valid conclusion based on the scatter plot alone.

```python
musketeers_5 = ...
```

```python
grader.check("q1_5")
```

<font color=red>**🚨 Important**: The tests in this section only check that you selected one of the possible answer choices, not that you have selected the right one!. Unlike in labs, tests in homeworks **do not** check that you answered correctly; they only check that your answer is *reasonable*, or in the correct format. To put it another way: all of your tests might pass, but that doesn't mean you'll get full credit – some of your answers may still be wrong. It's up to you to make sure that they're right!<font color=red>

## 2. Python Basics 🐍

**Question 2.1.** When you run the following cell, Python produces a cryptic error message.

```python
2027 = 2028 - 1.0
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `basics_1` below to indicate your answer.

1. The left hand side is an `int`, while the right hand side is a `float`. The left side should be `2027.0` instead.

1. The result should be written after the calculation. It should be `2028 - 1.0 = 2027`.

1. This is trying to create a variable named `2027`, which doesn't make sense because `2027` is a number.

1. Python is not able to subtract a `float` from an `int` because they are of different data types.


<font color=red>**🚨 Important**: Once you have finished this question, "comment"  out the above code cell out by replacing it with `# 2027 = 2028 - 1.0`. This will prevent the error message from appearing when your notebook is graded.</font>

***Note:*** A shortcut for "commenting" out code is to highlight the code and press `command` or `control` and `/`. This is especially useful if you want to comment many lines of code at once.

```python
basics_1 = ...
```

```python
grader.check("q2_1")
```

**Question 2.2.** Consider the following poorly-written code.

```python
two = 2
three = 3
three = three ** two
three = two - three
three = three * two
three = three * three
```

As this code executes, what values does the variable `three` take on? Assign 1, 2, 3, or 4 to `basics_2` to indicate your answer.

1. The variable `three` takes on the values 3, 6, 4, 8, 64.

1. The variable `three` takes on the values 3, 6, -4, -8, 64.

1. The variable `three` takes on the values 3, 9, 7, 14, 196.

1. The variable `three` takes on the values 3, 9, -7, -14, 196.

```python
basics_2 = ...
```

```python
grader.check("q2_2")
```

## 3. Road Trip   🚘 

This weekend, you plan on going on a road trip to Palm Springs with your friends Austin and Coleman. Before you go, you want to plan which route (Route A or Route B) and  whose car (Austin's or Coleman's) to take. Answer the questions below using Python to perform all the intermediate calculations such as adding, squaring, and dividing.

<font color=red>**🚨 Important**: The `math` package has not been imported. You don't need it for this question, and you should not import it, otherwise the Gradescope autograder may error.</font>

**Question 3.1.** First, you need to decide on the route. You plan to take the route for which the average speed is faster. 

For Route A, you will have to take 3 freeways, each with a different speed limit. Below is a table showing the speed limits (which we will assume is the speed you will travel at) and the time you will have to spend on each freeway on Route A.

| Freeway | Speed Limit (miles per hour) | Time (hours)|
| --- | --- | --- |
| I-5 | 70 | 3 |
| CA-73 | 55 | 2 |
| CA-133 | 45 | 1 |

Using this information, calculate the average speed, in miles per hour, if you take Route A, and assign your answer to the variable `route_A`. Recall from math and physics that the average speed is the total distance driven divided by the total time taken.

```python
# Feel free to define intermediate variables to use in your solution.
total_distance = ...
total_time = ...
route_A = ...
route_A
```

```python
grader.check("q3_1")
```

**Question 3.2.** Next, let's consider the other route you might take. For Route B, you will have to take 3 freeways, each with a different speed limit. Below is a table showing the speed limits (which we will assume is the speed you will travel at) and the *distance* you will need to travel on each freeway on Route B.

| Freeway | Speed Limit (miles per hour) | Distance (miles)|
| --- | --- | --- |
| I-5 | 70 | 150 |
| CA-57 | 65 | 110 |
| CA-109 | 55 | 90 |

Using this information, calculate the average speed, in miles per hour, if you take Route B, and assign your answer to the variable `route_B`.

Note that the third column is `'Distance (miles)'`, not `'Time (hours)'`. Unlike in Question 3.1, you aren't given the amount of time you'll spend on each freeway; you need to compute these times yourself. To calculate the time it will take on each freeway, divide the distance for that freeway by the speed for that freeway. Finally, add up the times for each freeway to find the total time.

```python
# Feel free to define intermediate variables to use in your solution.
total_distance2 = ...
total_time2 = ...

route_B = ...
route_B
```

```python
grader.check("q3_2")
```

**Question 3.3.** Now it's time to decide whose car to take (Austin's or Coleman's). You decide to take the car with the biggest fuel tank. 

Unfortunately, Austin doesn't know the exact volume of his car's fuel tank, only that the tank has a height of 16 inches, a width of 12 inches, and a length of 32 inches. Coleman doesn't know the exact volume of his car's fuel tank either, but knows that his tank is in the shape of a cube. What is the length of one of the sides of Coleman's fuel tank (in inches) so that it stores the same amount of fuel as Austin's? Save your answer in the variable `side_length`.

```python
# Feel free to define intermediate variables to use in your solution.
side_length = ...
side_length
```

```python
grader.check("q3_3")
```

In this problem, though you calculated three different quantities in three different ways, all of your results are actually considered **means**, of various kinds!

In Question 3.1, given $n$ values $x_1, x_2, ..., x_n$, you found an *arithmetic mean*, using the formula

$${x_1+x_2+...+x_n \over n},$$

where the numerator represented total distance and the denominator represented total time. An arithmetic mean is the usual type of mean or average you're used to seeing. It turns out that you actually computed a more sophisticated arithmetic mean, known as a _weighted arithmetic mean_, 

$$\frac{w_1 x_1 + w_2 x_2 + ... + w_n x_n}{w_1 + w_2 + ... + w_n},$$

where the weights $w_1, w_2, w_3$ were the times on each freeway.

In Question 3.2, given  $n$ values $x_1, x_2, ..., x_n$, you found a *harmonic mean*, using the formula

$${n \over {{1 \over x_1}+{1 \over x_2}+ ... + {1 \over x_n}}},$$ 


where the numerator represented total distance and the denominator represented total time. To calculate the total time, you needed to sum the time taken for each part of the trip, calculated using the fact that time is distance over speed. Again, it turns out that you actually computed the _weighted harmonic mean_, but this time the weights were the distances on each of the freeways. If you're curious, see the formula [here](https://en.wikipedia.org/wiki/Harmonic_mean#Weighted_harmonic_mean).

Finally in Question 3.3, given $n$ values $x_1, x_2, ..., x_n$, you found a *geometric mean*, using the formula 

$${\sqrt[n]{x_1 \cdot x_2 \cdot ... \cdot x_n}},$$ 

where each value represented a dimension of the fuel tank. 

As you can see, there are many different notions of the mean. You'll learn about some of them if you take DSC 40A!

## 4. AI Revolution ✨💻✨ 

In this problem, we want to compare and contrast early awareness of ChatGPT among American teenagers in different household income groups. The data below comes from [Pew Research Center's Survey of U.S. Teens](https://www.pewresearch.org/short-reads/2023/11/16/about-1-in-5-us-teens-whove-heard-of-chatgpt-have-used-it-for-schoolwork/sr_23-11-16_ai-in-schools_2-png/), in fall 2023, back when ChatGPT was new to the general public. The numbers below show **percentages** of each income group falling into each awareness category; note that each row sums to 100 or near 100 (because some respondents did not answer the question).

| Household Income     | Highly aware of ChatGPT  | A little aware of ChatGPT  | Not at all aware of ChatGPT |
|-------------------------------------------|-------------|-----------|-----------|
| less than 30,000 dollars      | 11 | 30 | 59 | 
| 30,000 to 74,999 dollars      | 22 | 36 | 40 | 
| more than 75,000 dollars   | 26 | 50  | 24 | 

We define the **dissimilarity** between two income groups as the largest absolute difference between their three respective percentages.

To better understand dissimilarity, consider the following hypothetical situation (not coming from the above table), where we compare the ChatGPT awareness between income group A and income group B. Suppose:
* Group A's *percentage of highly aware* is **10 percent more** than Group B's.
* Group A's *percentage of a little aware* is **4 percent less** than Group B's.
* Group A's *percentage of not at all aware* is **7 percent less** than Group B's.

Then, we would say the dissimilarity between Group A and Group B is 10, since 10 is larger than both 4 and 7.

**Question 4.1.** 
Using this method, compute the dissimilarity between the following two income groups: less than 30,000 dollars and 30,000 to 74,999 dollars.  Assign the result to the variable `dissimilarity`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting) rather than calculating the expression yourself. 

**_Hint:_**  The built-in `abs` function computes absolute values. 

```python
dissimilarity = ...
dissimilarity
```

```python
grader.check("q4_1")
```

**Question 4.2.** Which pair of income groups is **most** dissimilar, according to this measurement? Assign either 1, 2, or 3 to the variable `most_dissimilar` below. Check whether your answer matches up with your intuition.

1. Less than 30,000 dollars and 30,000 to 74,999 dollars. 
1. 30,000 to 74,999 dollars and more than 75,000 dollars.
1. Less than 30,000 dollars and more than 75,000 dollars.

```python
most_dissimilar = ...
```

```python
grader.check("q4_2")
```

**Question 4.3.** Suppose that instead of measuring awareness of ChatGPT with three categories (highly aware, a little aware, and not at all aware), the researchers had instead measured awareness with two categories (aware and not aware, where aware includes **both** highly aware and a little aware). Between which two income categories would the dissimilarity change if we measured awareness with two categories instead of three? Assign either 1, 2, 3, or 4 to the variable `would_change` below.

1. Less than 30,000 dollars and 30,000 to 74,999 dollars.
1. Less than 30,000 dollars and more than 75,000 dollars.
1. 30,000 to 74,999 dollars and more than 75,000 dollars.
1. All of the above.
1. None of the above.

```python
would_change = ...
```

```python
grader.check("q4_3")
```

## 5. Arrays 🗃️

**Question 5.1.** Make an array called `quirky_numbers` containing the following numbers (in the given order):

1. The square root of 82
2. 76 degrees, in radians
3. $4^3 + 6^5$
4. The mathematical constant $e$ divided by 2 (that is, $\frac{e}{2}$)
5. The base 10 logarithm of 34

**_Hint:_**  Check out the functions constants available in the `numpy` module, which has been imported as `np`. If you're unsure of what function to use, a quick Google search should do the trick.  Do **not** import `math` or any other modules. 

**_Note:_**  In this problem, as with all others, we'll only check that your answer is correct. There may be several valid ways to produce the correct answer.

```python
quirky_numbers = ...
quirky_numbers
```

```python
grader.check("q5_1")
```

**Question 5.2.** Make an array called `likes` containing the following three strings, in the specified order:
- `'I like eating'`
- `'art'`
- `'and bananas.'`



```python
likes = ...
likes
```

```python
grader.check("q5_2")
```

<center><img src=images/banana.jpg width=300><br><a href="https://www.nbcnews.com/news/us-news/man-who-ate-120-000-art-banana-said-he-would-n1098391">NYC artist who ate $120K banana off wall at Miami gallery says he would do it again</a></center>


In [Lecture 3](https://dsc10.com/resources/lectures/lec03/lec03.html#String-methods), we looked at several string methods, like `upper` and `replace`. Strings have another method that we haven't seen yet, called `join`. `join` takes one argument, an array of strings, and it returns a single string. Specifically, `some_string.join(some_array)` evaluates to a new string consisting of all of elements in `some_array`, with `some_string` inserted in between each element.

For example, `'-'.join(np.array(['call', '858', '534', '2230']))` evaluates to `'call-858-534-2230'`.

**Question 5.3.** Use the array `likes` and the method `join` to make two strings:

1. `'I like eating, art, and bananas.'` (call this one `by_comma`)
1. `'I like eating art and bananas.'` (call this one `by_space`)

```python
by_comma = ...
by_space = ...

# Don't change the lines below.
print(by_comma)
print(by_space)
```

```python
grader.check("q5_3")
```

Now let's get some practice accessing individual elements of arrays.  In Python (and in many programming languages), elements are accessed by *integer position*, with the position of the first element being zero. That's probably not the way you learned to count, so it's easy to get mixed up here. Be careful!

**Question 5.4.** The cell below creates an array of strings.

```python
some_strings = np.array(['sports', '⚽️', '🏀', '🏹', 'dog', '🐶', 'cat', '🐈'])
some_strings
```

What is the integer position of `'🏀'` in the array? You can just type in the answer, which should be of type `int`. This is a conceptual question, not a coding question.

**_Note:_**  Your answer should be a **positive** integer!

```python
basketball_position = ...
basketball_position
```

```python
grader.check("q5_4")
```

**Question 5.5.** Suppose you have an array with 350 elements. What is the integer position of the third-to-last element in this array? You can just type in the answer, which should be of type `int`. This is a conceptual question, not a coding question.

**_Note:_**  Again, your answer should be a **positive** integer!

```python
third_last_position = ...
third_last_position
```

```python
grader.check("q5_5")
```

**Question 5.6.** Suppose you have an array with 663 elements. At what integer position is the middle element of this array? You can just type in the answer, which should be of type `int`. This is a conceptual question, not a coding question.

**_Note:_**  Again, your answer should be a **positive** integer!

```python
mid_position = ...
mid_position
```

```python
grader.check("q5_6")
```

By the way, it's also possible to use negative integer positions to access elements in an array, which can be easier than using positive integer positions sometimes.  If a position is negative, you count from the end of the array rather than from the beginning. Position -1 corresponds to the last element, -2 corresponds to the second-to-last element, and so on. For instance, to find the third-to-last element of `some_strings`, we could use:

```python
some_strings[-3]
```

**Question 5.7.** Let's say you are trying create an online [multiplication table](https://en.wikipedia.org/wiki/Multiplication_table#Modern_times) for an elementary school. Using the function `np.arange()`, make an array called `multiples_of_10` which contains all the multiples of 10, in ascending order, that appear on the multiplication table below. 

It's also a good exercise to think about how your code would change if you wanted the multiples of $k$, for some other value of $k$ besides 10.

<center>
    <br>
    <img src=images/mult.jpg width=500>
</center>

```python
multiples_of_10 = ...
multiples_of_10
```

```python
grader.check("q5_7")
```

## 6. Unisex Names ⚥

The table below shows data on the eight most common unisex names in America. The data originally comes from the Social Security Administration, but it was compiled and curated by FiveThirtyEight for their 2015 article [The Most Common Unisex Names In America: Is Yours One Of Them?](https://fivethirtyeight.com/features/there-are-922-unisex-names-in-america-is-yours-one-of-them/).

For each name in the table below, we have the `'Count'`, which estimates the number of living people with that name, and the `'Male Share'`, or the proportion of people with that name who are male. Throughout this problem, we'll use *male* to mean assigned male at birth and *female* to mean assigned female at birth. We will assume that each person fits into exactly one of these two categories. 

|Name|Count|Male Share|
|---|---|---|
|Casey|177000|0.5843|
|Riley|155000|0.5076|
|Jessie|136000|0.4778|
|Jackie|133000|0.4211|
|Avery|122000|0.3352|
|Jaime|110000|0.5617|
|Peyton|95000|0.4337|
|Kerry|89000|0.4839|

In this question, we'll be working with the data from the `'Count'` and `'Male Share'` columns as *arrays* named `count` and `male_share`, respectively. Here are those arrays:

```python
count = np.array([177000, 155000, 136000, 133000, 122000, 110000, 95000, 89000])
count
```

```python
male_share = np.array([0.5843, 0.5076, 0.4778, 0.4211, 0.3352, 0.5617, 0.4337, 0.4839]) 
male_share
```

Remember, the `numpy` package (`np` for short) provides many handy functions for working with arrays. These are specifically designed to work with arrays and are faster than using Python's built-in functions. 

Some frequently used array functions are `np.min()`, `np.max()`, `np.sum()`, `np.abs()`, and `np.round()`. There are many more, which you can browse by typing `np.` into a code cell and hitting the *tab* key.

**Question 6.1.** What proportion of people with each name are female? Store the female share proportions for each name in a new array called `female_share`. 

```python
female_share = ...
female_share
```

```python
grader.check("q6_1")
```

**Question 6.2.** Now let's find the gap between the male and female shares for each name. Use `male_share` and `female_share` to create an array called `gaps` containing the absolute differences between the male share and female share for each name.

```python
gaps = ...
gaps
```

```python
grader.check("q6_2")
```

**Question 6.3.** Now, find the gaps between the male and female shares for each name, but this time, you're only allowed to use the variable `male_share`. You may not use `female_share`. Create an array called `gaps_again` containing the absolute differences between the male share and female share for each name. The answer will be the same as the last question, but your code should be different.

```python
gaps_again = ...
gaps_again
```

```python
grader.check("q6_3")
```

**Question 6.4.** You might say that the least unisex name is the one with the largest gap. Find the largest value in the `gaps` array and save it as `largest_gap`. Referring back to the table, try to figure out which name that is!

```python
largest_gap = ...
largest_gap
```

```python
grader.check("q6_4")
```

**Question 6.5.** Recall that the `count` array contains the estimated number of people with each name. Create an array called `male_count` containing the estimated number of males with each name, and another called `female_count` containing the estimated number of females with each name. Make sure to round each value to the nearest whole number.

```python
male_count = ...
female_count = ...
print('male_count: ', male_count, '\nfemale_count: ', female_count) # Don't change this line!
```

```python
grader.check("q6_5")
```

**Question 6.6.** Let's try to figure out the most unisex name as follows. Assign `smallest_diff` to the smallest absolute difference between the number of males and number of females with a certain name, among the given names. Referring back to the table, try to figure out which name that is!

```python
smallest_diff = ...
smallest_diff
```

```python
grader.check("q6_6")
```

## 7. Disney Movies 👸🏰🌠

In this question, we will explore movies produced by *Walt Disney Studios* from 1937 to 2016.

<img src="images/walt-disney-logo.webp" width=60%>

The file `disney_movies.csv` in the `data/` directory contains information about 579 such Disney movies. Its columns are described below.

| Column      | Description |
| ----------- | ----------- |
| `'movie_title'`      | Name of the movie     |
| `'genre'` | Genre of the movie  |
| `'MPAA_rating'` | Content rating, determining suitability for children |
| `'total_gross'` | Total box office revenue |
| `'adjusted_gross'` | Total box office revenue, after adjustment for inflation |
| `'year'`   | Year of release   |

**Question 7.1.** Read this file into a DataFrame called `disney_movies`. 

```python
disney_movies = ...
disney_movies
```

```python
grader.check("q7_1")
```

**Question 7.2.** Since the values in `total_gross` are so large, let's add a column to `disney_movies` called `'total_gross_millions'` that contains the total box office revenue **in millions of dollars**. Similarly, create a column called `'adjusted_gross_millions'` with the inflation-adjusted amounts in millions of dollars. Then, drop the original `'total_gross'` and `'adjusted_gross'` columns. Do not round.

```python
disney_movies = ...
disney_movies
```

```python
grader.check("q7_2")
```

**Question 7.3.** Create a new DataFrame, called `disney`, by setting the index of `disney_movies` to `'movie_title'`. Don't change `disney_movies`.

```python
disney = ...
disney
```

```python
grader.check("q7_3")
```

You should think about why we've chosen to set the index to `'movie_title'`, instead of any other column.

**Question 7.4.** Time for some movie trivia! What genre is the movie `'Frozen'`❄ and in what year was it released? Assign your answers to `frozen_genre`, which should be a string, and `frozen_year`, which should be an `int`.

Don't type in the answers by hand; get Python to extract this information for you.

```python
frozen_genre = ...
frozen_year = ...

# Don't change the lines below.
print('Genre', frozen_genre)
print('Year:', frozen_year)
```

```python
grader.check("q7_4")
```

**Question 7.5.** Disney has produced many successful films over the years. Find the number of movies with an adjusted box office revenue of over 200 million dollars. Assign your answer to `disney_200m`.

```python
disney_200m = ...
disney_200m
```

```python
grader.check("q7_5")
```

**Question 7.6.** Assign `third_highest_revenue` to the third-highest value of `adjusted_gross_millions`. Assign `third_highest_name` to the name of the movie with the third-highest adjusted revenue.

Again, don't type in these values by hand; get Python to extract this information for you.

**_Note:_** Remember that you can perform intermediate steps by adding new cells above the provided ones, or by adding in additional lines of code.

```python
third_highest_revenue = ...
third_highest_name = ...

# Don't change the lines below.
print('Revenue:', third_highest_revenue, 'million dollars')
print('Name:', third_highest_name)
```

```python
grader.check("q7_6")
```

**Question 7.7.** Disney is known for making movies for children, many of which are rated G for "general audiences". Let's see how much money Disney has earned in total across all of its G-rated movies. Find the sum of the `total_gross_millions` column for G-rated movies only, and assign your answer to `g_revenue`.

```python
g_revenue = ...
g_revenue
```

```python
grader.check("q7_7")
```

**Question 7.8.** Finally, let's see which genres are most popular by seeing how much money movies from each genre earn on average. Create a DataFrame called `revenue_by_genre` that is indexed by `genre` and contains just one column, called `adjusted_gross_millions`, which contains the average adjusted box office revenue of all movies from a given genre. Order the rows of `revenue_by_genre` in descending order of `adjusted_gross_millions` to see what genres Disney does best! 🎵

**_Hint:_** Our solution for this question used only one line of code (thanks, `groupby`)!

```python
revenue_by_genre = ...
revenue_by_genre
```

```python
grader.check("q7_8")
```

## Finish Line: Almost there, but make sure to follow the steps below to submit! 🏁

<hr style="color:Maroon;background-color:Maroon;border:0 none; height: 3px;">

**Citations:** Did you use any generative artificial intelligence tools to assist you on this assignment? If so, please tell us which tools you used by setting `gen_ai` to a **list** of numbers. Include all tools used, from the options below:

1. AI Tutor, in Tutor Mode
2. AI Tutor, in ChatGPT Mode
3. Any other outside generative AI tool, such as Gemini, Claude, ChatGPT, etc.

If you didn't use any such tools, set `gen_ai` to an empty list, `[]`.

```python
gen_ai = ...
```

```python
grader.check("citations")
```

<hr style="color:Maroon;background-color:Maroon;border:0 none; height: 3px;">

To submit your assignment:

1. Select `Kernel -> Restart Kernel and Run All Cells` to ensure that you have executed all cells, including the test cells.
2. Read through the notebook to make sure everything is fine and all tests passed.
3. Run the cell below to run all tests, and make sure that they all pass.
4. Download your notebook as a `.ipynb` file using `File -> Download`, then upload your notebook  to Gradescope under "Lab 0".
5. **Stick around while the Gradescope autograder grades your work. Make sure you see that all tests have passed on Gradescope.**
6. Check that you have a confirmation email from Gradescope and save it as proof of your submission.

With homeworks, unlike with labs, the grade you see on Gradescope is **not your final score**. We will run correctness tests after the assignment's due date has passed.

```python
grader.check_all()
```

Tests for citations (points: None)
```
>>> isinstance(gen_ai, list)
True
```
failure_message: In Python, a list has square brackets and comma-separated entries, such as [1, 2, 3]


Tests for q1_1 (points: None)
```
>>> musketeers_1 in [1, 2, 3, 4, 5]
True
```


Tests for q1_2 (points: None)
```
>>> musketeers_2 in [1, 2, 3, 4, 5]
True
```


Tests for q1_3 (points: None)
```
>>> musketeers_3 in [1, 2, 3, 4, 5]
True
```


Tests for q1_4 (points: None)
```
>>> musketeers_4 in [1, 2]
True
```


Tests for q1_5 (points: None)
```
>>> musketeers_5 in [1, 2, 3, 4, 5]
True
```


Tests for q2_1 (points: None)
```
>>> import numbers
>>> isinstance(basics_1, numbers.Integral)
True
```
```
>>> basics_1 in [1, 2, 3, 4]
True
```


Tests for q2_2 (points: None)
```
>>> import numbers
>>> isinstance(basics_2, numbers.Integral)
True
```
```
>>> basics_2 in [1, 2, 3, 4]
True
```


Tests for q3_1 (points: None)
```
>>> route_A > 0
True
```
```
>>> int(route_A * 10 ** 4 % 10 ** 4) == 8333
True
```
failure_message: Try again!


Tests for q3_2 (points: None)
```
>>> route_B > 0
True
```
```
>>> int(route_B * 10 ** 10 % 10 ** 6) == 4564
True
```
failure_message: Try again!


Tests for q3_3 (points: None)
```
>>> side_length > 0
True
```
```
>>> int(side_length * 10 ** 7 % 10 ** 3) == 278
True
```
failure_message: Try again!


Tests for q4_1 (points: None)
```
>>> 0 < dissimilarity < 100
True
```
```
>>> int(dissimilarity ** 0.5 * 7 ** 2 % 10 ** 3) == 213
True
```
failure_message: Try again!


Tests for q4_2 (points: None)
```
>>> import numbers
>>> isinstance(most_dissimilar, numbers.Integral)
True
```
```
>>> most_dissimilar in [1, 2, 3]
True
```


Tests for q4_3 (points: None)
```
>>> import numbers
>>> isinstance(would_change, numbers.Integral)
True
```
```
>>> would_change in [1, 2, 3, 4, 5]
True
```


Tests for q5_1 (points: None)
```
>>> isinstance(quirky_numbers, np.ndarray) and quirky_numbers.size == 5
True
```
```
>>> int(quirky_numbers[0] * 10 ** 4 % 10 ** 3) == 553
True
```
failure_message: Something is wrong with the first element. Try again!
```
>>> int(quirky_numbers[1] * 10 ** 3 % 10 ** 3) == 326
True
```
failure_message: Something is wrong with the second element. Try again!
```
>>> quirky_numbers[2] % 10 ** 3 / 10 == 84
True
```
failure_message: Something is wrong with the third element. Try again!
```
>>> int(quirky_numbers[3] * 10 ** 3 % 10 ** 3) == 359
True
```
failure_message: Something is wrong with the fourth element. Try again!
```
>>> int(quirky_numbers[4] * 10 ** 3 % 10 ** 3) == 531
True
```
failure_message: Something is wrong with the last element. Try again!


Tests for q5_2 (points: None)
```
>>> isinstance(likes, np.ndarray)
True
```
```
>>> likes[0] == 'I like eating'
True
```


Tests for q5_3 (points: None)
```
>>> isinstance(by_comma, str) and isinstance(by_space, str)
True
```
```
>>> len(by_comma) == 32
True
```
failure_message: You have the wrong length for by_comma. Make sure by_comma appears EXACTLY as described.


Tests for q5_4 (points: None)
```
>>> import numbers
>>> isinstance(basketball_position, numbers.Integral) and basketball_position > 0
True
```
failure_message: Use a positive integer position.


Tests for q5_5 (points: None)
```
>>> import numbers
>>> isinstance(third_last_position, numbers.Integral) and third_last_position > 0
True
```
failure_message: Use a positive integer position.


Tests for q5_6 (points: None)
```
>>> import numbers
>>> isinstance(mid_position, numbers.Integral) and mid_position > 0
True
```
failure_message: Use a positive integer position.


Tests for q5_7 (points: None)
```
>>> isinstance(multiples_of_10, np.ndarray) and len(multiples_of_10) == 12
True
```
failure_message: Refer to the multiplication table linked in the question.


Tests for q6_1 (points: None)
```
>>> isinstance(female_share, np.ndarray)
True
```
```
>>> int(female_share[2] * 10 ** 3 % 10 ** 2) == 22
True
```
failure_message: Try again!
```
>>> int(female_share[6] * 10 ** 3 % 10 ** 2) == 66
True
```
failure_message: Try again!


Tests for q6_2 (points: None)
```
>>> isinstance(gaps, np.ndarray)
True
```
```
>>> int(gaps[1] * 10 ** 3 % 10 ** 2) == 15
True
```
failure_message: Try again!
```
>>> int(gaps[7] * 10 ** 3 % 10 ** 2) == 32
True
```
failure_message: Try again!


Tests for q6_3 (points: None)
```
>>> isinstance(gaps_again, np.ndarray)
True
```
```
>>> int(gaps_again[1] * 10 ** 3 % 10 ** 2) == 15
True
```
failure_message: Try again!
```
>>> int(gaps_again[7] * 10 ** 3 % 10 ** 2) == 32
True
```
failure_message: Try again!


Tests for q6_4 (points: None)
```
>>> 0 <= largest_gap <= 0.5
True
```
```
>>> int(largest_gap * 10 ** 3 % 10 ** 2) == 29
True
```
failure_message: Try again!


Tests for q6_5 (points: None)
```
>>> isinstance(male_count, np.ndarray) and isinstance(female_count, np.ndarray)
True
```
```
>>> int(male_count[2] / 10 % 10 ** 3) == 498
True
```
failure_message: Try again!
```
>>> int(male_count[5] / 10 % 10 ** 3) == 178
True
```
failure_message: Try again!
```
>>> int(female_count[2] / 10 % 10 ** 3) == 101
True
```
failure_message: Try again!
```
>>> int(female_count[4] / 10 % 10 ** 3) == 110
True
```
failure_message: Try again!


Tests for q6_6 (points: None)
```
>>> 1000 < smallest_diff < 100000
True
```
```
>>> int(smallest_diff / 8 % 10 ** 3) == 294
True
```
failure_message: Try again!


Tests for q7_1 (points: None)
```
>>> isinstance(disney_movies, bpd.DataFrame)
True
```


Tests for q7_2 (points: None)
```
>>> 'total_gross_millions' in disney_movies.columns and 'adjusted_gross_millions' in disney_movies.columns
True
```
```
>>> len(disney_movies.columns) == 6
True
```
```
>>> int(disney_movies.get('total_gross_millions').iloc[0] * 10 ** 3 % 10 ** 2) == 25
True
```
failure_message: Try again!
```
>>> int(disney_movies.get('adjusted_gross_millions').iloc[0] * 10 ** 3 % 10 ** 2) == 53
True
```
failure_message: Try again!


Tests for q7_3 (points: None)
```
>>> disney.shape == (579, 5)
True
```


Tests for q7_4 (points: None)
```
>>> import numbers
>>> isinstance(frozen_genre, str) and isinstance(frozen_year, numbers.Integral)
True
```


Tests for q7_5 (points: None)
```
>>> import numbers
>>> isinstance(disney_200m, numbers.Integral)
True
```


Tests for q7_6 (points: None)
```
>>> isinstance(third_highest_revenue, float) and isinstance(third_highest_name, str)
True
```
```
>>> int(third_highest_revenue * 10 ** 3 % 10 ** 2 * 10) == 908
True
```
failure_message: Try again!


Tests for q7_7 (points: None)
```
>>> isinstance(g_revenue, float)
True
```
```
>>> int(g_revenue * 10 ** 3 % 10 ** 2) == 92
True
```
failure_message: Try again!


Tests for q7_8 (points: None)
```
>>> isinstance(revenue_by_genre, bpd.DataFrame) and revenue_by_genre.shape == (12, 1)
True
```
```
>>> 'Adventure' in revenue_by_genre.index
True
```
```
>>> int(revenue_by_genre.get('adjusted_gross_millions').iloc[4] * 10 ** 3 % 10 ** 2) == 67
True
```
failure_message: Try again!
